<table>
  <tr>
    <td><center><img src="../../img/mlhep-logo-transparent-400.png" width="400"></center></td>
    <td><h1><center>The Sixth Machine Learning <br/>for High Energy Physics Summer School 2020</center></h1></td>
  </tr>
 </table>

## Introduction to the seminar structure
This notebook serves as a very basic part that should help you get familiar with different kinds of tasks you can throughout the school seminars. From easiest to more and more difficult those range from:
- *follow the speaker* or just regular python cells that one can execute inside own `cocalc` project environment.
- *write your own piece of code*
- *test your code with explicit tests* and ask for assistance if it doesn't work
- *be ready to run your code through external tests*

So, having those shades of gray in mind, we start from a simple 2000-years old algorithm, then peek at system resources available. Afterwards, we accomplish code writing assignment and move on to the tests that help you getting scored.

### Simple function for finding greatest common denominator

```python
a = int(input("Number A"))
```

Output:
```
Number A 2134
```

```python
b = int(input("Number B"))
```

Output:
```
Number B 21341234
```

```python
def gcd(a, b):
    """ Euclidean algorithm for finding greatest common denominator 
    a, b - integer numbers
    """
    if a < b:
        return gcd(b, a)
    if a % b > 0:
        return gcd(a, a % b)
    else:
        return b
```

**Hint**: Press `Shift-tab` inside function argument brackets of the next cell to see the help snippet.

```python
gcd(a, b)
```

Output:
```
2
```

### System resources check

```python
# path
!pwd
```

Output:
```
/home/user/mlhep2020/Section_1/1.1-intro
```

```python
# free memory
!free -m
```

Output:
```
total        used        free      shared  buff/cache   available
Mem:         225680        3617      200363           1       21699      220556
Swap:             0           0           0
```

```python
# disk resources
!df -h
```

Output:
```
Filesystem                                        Size  Used Avail Use% Mounted on
overlay                                           124G   33G   91G  27% /
tmpfs                                              64M     0   64M   0% /dev
tmpfs                                             111G     0  111G   0% /sys/fs/cgroup
10.0.57.35:/e43509a3-d19b-4dec-bdfe-e4e6c53e5648  124G   36G   88G  30% /home/user
/dev/sda1                                         124G   33G   91G  27% /etc/hosts
shm                                                64M     0   64M   0% /dev/shm
tmpfs                                             111G   12K  111G   1% /proc/driver/nvidia
udev                                              111G     0  111G   0% /dev/nvidia0
tmpfs                                             111G     0  111G   0% /proc/acpi
tmpfs                                             111G     0  111G   0% /proc/scsi
tmpfs                                             111G     0  111G   0% /sys/f
```

### Let's write some code and test it

<!-- grade_id: a47442, points: 0, grade: false, solution: true -->
```python
def sum_1_to_n(N):
    """
    Given integer N compute the sum of integer numbers from 0 to N
    if N is less than zero, return -sum_1_to_n(-N)
    """
    if N >= 0:
        result = N * (N+1) / 2
    else:
        return -sum_1_to_n(-N)
    return result
```

```python
sum_1_to_n(10)
```

Output:
```
55.0
```

```python
from nose.tools import assert_equal, timed
import numpy.testing as np_testing
import time
```

**Explicit local tests**

Those test can verify different aspects of your code like accuracy, timing or anything one can put inside an assert statement of similar kind.
If any of those do not pass, and you feel like you got stuck, it is good time to ask for **assistance**.

<!-- grade_id: simple_checks, points: 1, grade: true, solution: false -->
```python
np_testing.assert_equal(sum_1_to_n(10), 55)
np_testing.assert_equal(sum_1_to_n(0), 0)
```

**External tests to be checked on instructor side**

Those are essentially doing the same thing as explicit local ones, but the verification happens on instructor's side, so you cannot execute the following two cells on your own.

<!-- grade_id: accuracy_check, points: 1, grade: true, solution: false -->
```python
# This cell is going to check function output
```

<!-- grade_id: speed_check, points: 1, grade: true, solution: false -->
```python
# This cell is going to check function speed

```

## That's it!
If all explicit test above has finished successfully, you know how to address typical code problems in a meaningful way. Now you have to be patient before getting scored results since processing tests on the instructor side takes some time.